```python
import torch
import torch.nn as nn
```

## Problem 1

Write a function that computes the attention between query, key and value vectors.\
The attention formula is $Attention(Q,K,V) = Softmax(\frac{QK^T}{\sqrt{d_k}}) V$ where $d_k$ is the hidden dimension of the key vectors.

<!-- grade_id: cell-a4b06670d7392040, points: 0, grade: false, solution: true -->
```python
def attention(q, k, v):
    """Compute the attention between Q, K and V
    All tensors have shape (batch size, sequence length, hidden dim)
    """
    # YOUR CODE HERE
    raise NotImplementedError()
```

<!-- grade_id: cell-d769f61e3fc66dfe, points: 5, grade: true, solution: false -->
```python
q = torch.zeros(4, 16, 32)
k = torch.zeros(4, 16, 32)
v = torch.zeros(4, 16, 32)

expected = torch.zeros(4, 16, 32)
assert torch.allclose(expected, attention(q, k, v))
```

## Problem 2

Write a custom layer that implements the following formula:

$h_1 = W_1 x + b_1$\
$h_2 = W_2 h_1 + b_2$\
$h_{3} = h_1 + h_2$\
$y = \text{sigmoid}(W_{\text{3}} h_{3} + b_{\text{3}})$

where every vector is of size 10 and $W_1, b_1, W_2, b_2, W_3, b_3$ are learnable parameters.


<!-- grade_id: cell-0ba30c835d03c24c, points: 0, grade: false, solution: true -->
```python
class MyModel(nn.Module):
    def __init__(self):
        super(MyModel, self).__init__()
        # YOUR CODE HERE
        raise NotImplementedError()

    def forward(self, x):
        # YOUR CODE HERE
        raise NotImplementedError()

```

<!-- grade_id: cell-a77e13b85233be8c, points: 5, grade: true, solution: false -->
```python
model = MyModel()
x = torch.randn(4, 10)
y = model(x)
assert y.shape == (4, 10)
```

## Problem 3

The MNIST dataset is a collection of handwritten digits. The goal is to classify each digit into one of 10 classes.
Each image is 28x28 pixels, and each pixel is a grayscale value between 0 and 255. There are 60000 training images and 10000 test images.

Train a neural network that achieves > 90% accuracy on the MNIST dataset.\
Hint: Convolutional Networks are effective for image classification tasks.

```python
# Load MNIST dataset
from torchvision import datasets, transforms


# Load training data
train_dataset = datasets.MNIST(root='./data', train=True, download=True, transform=transforms.ToTensor())
train_loader = torch.utils.data.DataLoader(train_dataset, batch_size=64, shuffle=True)

# Load test data
test_dataset = datasets.MNIST(root='./data', train=False, download=True, transform=transforms.ToTensor())
test_loader = torch.utils.data.DataLoader(test_dataset, batch_size=1000, shuffle=False)

print(f"Training dataset size: {len(train_dataset)}")
print(f"Test dataset size: {len(test_dataset)}")

def evaluate(model):
    model.eval()
    correct = 0
    total = 0
    with torch.no_grad():
        for data, target in test_loader:
            output = model(data)
            pred = output.argmax(dim=1, keepdim=True)
            correct += pred.eq(target.view_as(pred)).sum().item()
            total += target.size(0)

    return correct / total
```

<!-- grade_id: cell-f248b5896dbf97d5, points: 0, grade: false, solution: true -->
```python
# Create your model in the "model" variable and train it
# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: cell-c6c1ef1811be9dc3, points: 10, grade: true, solution: false -->
```python
x = torch.randn(4, 1, 28, 28)
# Assert forward pass works
y = model(x)
assert y.shape == (4, 10)
assert evaluate(model) > 0.9
```